rule_files:
  - ../../monitoring/base/victoriametrics/rules/converted/monitoring-alertrule.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'up{job="victoriametrics-operator"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: VMOperatorDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: VictoriaMetrics operator has disappeared.
  - interval: 1m
    input_series:
      - series: 'up{job="vmalertmanager"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: VMAlertmanagerDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: VMAlertmanager has disappeared.

  - interval: 1m
    input_series:
      - series: 'up{job="vmagent-smallset"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: VMAgentSmallsetDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: VMAgent for smallset has disappeared.
  - interval: 1m
    input_series:
      - series: 'up{job="vmalert-smallset"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: VMAlertSmallsetDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: VMAlert for smallset has disappeared.
  - interval: 1m
    input_series:
      - series: 'up{job="vmsingle-smallset"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: VMSingleSmallsetDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: VMSingle for smallset has disappeared.

  - interval: 1m
    input_series:
      - series: 'up{job="vmagent-largeset"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: VMAgentLargesetDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: VMAgent for largeset has disappeared.
  - interval: 1m
    input_series:
      - series: 'up{job="vmalert-largeset"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: VMAlertLargesetDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: VMAlert for largeset has disappeared.
  - interval: 1m
    input_series:
      - series: 'up{job="vmstorage-largeset"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: VMStorageLargesetDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: VMStorage for largeset has disappeared.
  - interval: 1m
    input_series:
      - series: 'up{job="vmselect-largeset"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: VMSelectLargesetDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: VMSelect for largeset has disappeared.
  - interval: 1m
    input_series:
      - series: 'up{job="vminsert-largeset"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: VMInsertLargesetDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: VMInsert for largeset has disappeared.

  - interval: 1m
    input_series:
      - series: 'up{job="alertmanager"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: AlertmanagerDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: Alertmanager has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="prometheus"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: PrometheusDown
        exp_alerts:
          - exp_labels:
              severity: minor
            exp_annotations:
              runbook: TBD
              summary: Prometheus has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="pushgateway"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: PushGatewayDown
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              runbook: TBD
              summary: pushgateway is down.
